# Julia demo assignment

<!-- grade_id: cell-43b5e387cd8c1f44, points: 0, grade: false, solution: true -->
```julia
# Write a function to return the volume of a sphere
function sphere_vol(r)
    ### BEGIN SOLUTION
    return 4/3*pi*r^3
    ### END SOLUTION
end
```

```julia
# Should be 113.097...
sphere_vol(3)

```

<!-- grade_id: cell-1311828b8e7870fa, points: 1, grade: true, solution: false -->
```julia
@assert floor(sphere_vol(3)) == 113.0
### BEGIN HIDDEN TESTS
@assert floor(sphere_vol(4)) == 268
### END HIDDEN TESTS
```

<!-- grade_id: cell-c415c19849c572a7, points: 0, grade: false, solution: true -->
```julia
# Write a function to calculate x for 0 = a*x^2+b*x+c
function quadratic(a::Float64, b::Float64, c::Float64)
    ### BEGIN SOLUTION
    # nested function
    quad(a, sqr_term, b) = (-b + sqr_term) / 2a
    
    # unlike other languages 2a is equivalent to 2*a
    # a^2 is used instead of a**2 or pow(a,2)
    sqr_term = sqrt(b^2-4a*c)
    r1 = quad(a, sqr_term, b)
    r2 = quad(a, -sqr_term, b)
    # multiple values can be returned from a function using tuples
    # if the return keyword is omitted, the last term is returned
    r1, r2
    ### END SOLUTION
end
```

```julia
quad1, quad2 = quadratic(2.0, -2.0, -12.0)
```

<!-- grade_id: cell-69edfc8550623b53, points: 1, grade: true, solution: false -->
```julia
quad1, quad2 = quadratic(2.0, -2.0, -12.0)
@assert quad1 == 3.0
@assert quad2 == -2.0
### BEGIN HIDDEN TESTS
quad3, quad4 = quadratic(1.0, -4.0, -23.0)
@assert floor(quad3) == 7.0
@assert floor(quad4) == -4.0
### END HIDDEN TESTS
```

```julia

```