## This Notebook - Goals - FOR EDINA

**What?:**
- What is the purpose of this notebook?

**Who?:**
- Who is this aimed at? Who is likely to come across it?

**Why?:**
- What is the motivation behind this notebook?

**Noteable features to exploit:**
- what features of noteable are being exploited?

**How?:**
- How are the features exploited? 
- What is being done to satisfy the needs of the audience?

<hr>

<div class="alert alert-info" role="alert">

# Notes for teachers

If you are learning to use nbgrader, this notebook is for you.

This notebook demonstrates the particular syntax used in autograded assignments.

There is an [accompanying guide] that explains the process of uploading this to formgrader (the toolbar used by nbgrader).

If you would like a worksheet template please see [THIS NOTEBOOK] which you can download and edit to suit your assignment.

## How to use this sheet

To see the markdown used in any markdown cell, click on it and press <code>Enter</code>.

</div>

# Introduction to Python - Class Test

This test will ask you questions on a few of the python topics studied in your introduction to python class. Some of the questions in the test are automarked - take note of the tests you can see, and be aware your grade will also be based on hidden tests, only visible to the instructor.

## Question 1 (5 marks)

This question requires a response in LaTeX:

Write the formula for the circumference of a circle in terms of radius.

<!-- grade_id: circumference, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2 (10 marks)

Here is a question that requires a code response:

Given your answer above, write a python function in the cell below that calculates the circumference of a circle, given the radius.

Function specification:
- Function name - cirm
- Input - radius of circle
- Output - circumference of circle, correct to 2 decimal places

<i>Hint: You will need to import the 'math' module in order to use $\pi$ in your calculation.</i>

<!-- grade_id: cirm, points: 0, grade: false, solution: true -->
```python
# import required libraries here

# function cirm
def cirm(n):
    
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: tests_cirm, points: 1, grade: true, solution: false -->
```python
assert round(cirm(1),2) == 6.28
```

<!-- grade_id: tests_cirm_hidden, points: 6, grade: true, solution: false -->
```python
assert round(cirm(2),2) == 12.57

```

Write a function that calculates a list of circumferences, given a list of radii.

Function specification:
- Function name - cirms
- Input - list of radii: (eg: <code>\[1, 2, 3, 4, 5\]</code> )
- Output - list of associated circumferences, rounded to 2 decimal places

<!-- grade_id: cirms, points: 0, grade: false, solution: true -->
```python
def cirms(ns):
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

Now use your function to calculate the circumference of circles with radii:

\[1, 2, 3, 4, 5\]

You should store the list of circumferences in a variable, <code>circumferences</code>.

<!-- grade_id: first5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: test_cirms, points: 3, grade: true, solution: false -->
```python
assert cirms([1,2]) == [6.28, 12.57]

```

## Question 3 (1 mark)

In a maths lesson, Bethany is asked to guess the circumference of the circles. Here are her guesses:

| Radius | Guess |
|--------|-------|
|    1   |   6   |
|    2   |  14   |
|    3   |  20   |
|    4   |  25   |
|    5   |  35   |

<br>

In the <b>markdown cell</b> below, create a table similar to that above, also including the actual values (as calculated in question 2).

<!-- grade_id: table, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4 (9 marks)

Produce a suitable plot for Bethany's guesses. Again, consider suitable library functions to achieve this. 

<i><b>Hint:</b> recall that you have a list of true values, and a list of guessed values is given in the cell below.</i>

```python
# this code cell stores Bethany's guesses in a list called gs
gs = [6, 14, 20, 25, 35]
```

<!-- grade_id: bethanyPlot1, points: 9, grade: true, solution: true -->
```python
# import helpful libraries here

# YOUR CODE HERE
raise NotImplementedError()
```

## Question 5 (5 marks)

Plot the graph from question 4 again, but this time include a line showing the true circumference values.

_Hint: you can copy and paste cells by using the toolbar above "cut, copy, paste"._

<!-- grade_id: bethanyPlot2, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 6 (15 marks)

Reflect on the tools used in this notebook. Use as many code/markdown cells as you like to comment on the quality of Bethany's guesses. Consider calculating some measure of error.


<hr>

### Submission instructions

Once you have completed the tasks above, please submit your assignment on Noteable.

Multiple submissions are allowed - the latest will be counted.